---
title: 'Clean your data in a reproducible way'
description: 'Fix potential issues in the dataset'
---

## Insert exercise title here

```yaml
type: NormalExercise
key: ab647ffeb9
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
