---
title: 'Disclosure Risk Treatment and Anonymisation'
description: 'Prepare your microdata for anonymisation'
---

## Set up identifiers, key & sentitive variables

```yaml
type: NormalExercise
key: 2e0d16dd49
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
