---
title: 'Deal with hierarchical data'
description: 'Connect nested frame'
---

## Set up primary key

```yaml
type: PureMultipleChoiceExercise
key: a932f4f670
xp: 50
```

What type of sample can you have to deal with?

`@hint`


`@possible_answers`
simple random
cluster
Stratified

`@feedback`


---

## Link nested frame

```yaml
type: NormalExercise
key: d10c603df5
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
