---
title: 'Calculated indicators'
description: 'Create new variables from the existing ones'
---

## Dummy variable in the same frame

```yaml
type: NormalExercise
key: b09e494ed5
xp: 100
```

Dummy variable in the same frame

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Dummy variable in nested frame

```yaml
type: NormalExercise
key: b4fed8d7bf
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
