---
title: 'Clean your data in a reproducible way'
description: 'Fix potential issues in the dataset'
---

## Recoding modalities

```yaml
type: NormalExercise
key: ab647ffeb9
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Cleaning from or_other

```yaml
type: NormalExercise
key: 34522eaafe
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Dropping numeric outliers

```yaml
type: NormalExercise
key: 10ee294dcc
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
