---
title: 'Configure project'
description: 'Configure path and sampling plan'
---

## Introduction

```yaml
type: VideoExercise
key: 8d01a9d567
xp: 50
```

`@projector_key`
670bc8dc3b01792475f2a861549f611c



## An exercise title written in sentence case

```yaml
type: NormalExercise
key: c909bf2ad2
lang: r
xp: 100
skills: 1
```

You have just collected a complex survey data set using Kobotoolobox, ODK or ONA. You need now to do the analysis. In this first chapter, we will configure your project so that the analysis is properly organized and reproducible for your peer.

---

## Initiate the project

```yaml
type: NormalExercise
key: 9af498a604
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Configure path

```yaml
type: NormalExercise
key: aaef2fcc5f
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
