---
title: 'Configure project'
description: 'Configure path and sampling plan'
---

## Introduction

```yaml
type: VideoExercise
key: 8d01a9d567
xp: 50
```

`@projector_key`
670bc8dc3b01792475f2a861549f611c

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 6f03a9a4ee
xp: 50
```

`@projector_key`
6288eb291de5f36a436898290a33e2c2


---

## An exercise title written in sentence case

```yaml
type: NormalExercise
key: c909bf2ad2
lang: r
xp: 100
skills: 1
```

---

## Initiate the project

```yaml
type: NormalExercise
key: 9af498a604
xp: 100
```



`@instructions`

First start a new project within your Rstduio Environment. 

Now load the koboloadeR package. 

The very first step is to Initiate a project. This will allwo to create a series of standard folders within the project folder.



`@hint`


`@pre_exercise_code`
```{r}
library(koboloadeR)
kobo_project_init()

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Configure path

```yaml
type: NormalExercise
key: aaef2fcc5f
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---


