---
title: 'Humanitarian Survey'
description: 'Common practice for survey data collection in humanitarian settings'
---

## An exercise title written in sentence case

```yaml
type: NormalExercise
key: c909bf2ad2
lang: r
xp: 100
skills: 1
```

You have just collected a complex survey data set using Kobotoolobox, ODK or ONA. You need now to do the analysis. In this first chapter, we will configure your project so that the analysis is properly organized and reproducible for your peer.

`@instructions`
- Load the koboloaedeR package
- Initialize the project 
- Set up the path to your data and form

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`
```{r}
# Load datasets and packages here.
```

`@sample_code`
```{r}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`
```{r}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```
